# Deep Learning

In this exercise, you will use a deep neural network to predict the values of houses based on some provided input data. You will use keras to build the model. Below is a description of how the keras models are set up.

**Please read more about the data [here](https://www.cs.toronto.edu/~delve/data/boston/bostonDetail.html).**

```python
!pip install --upgrade tensorflow==2.0
```

Output:
```
Collecting tensorflow==2.0
[?25l  Downloading https://files.pythonhosted.org/packages/46/0f/7bd55361168bb32796b360ad15a25de6966c9c1beb58a8e30c01c8279862/tensorflow-2.0.0-cp36-cp36m-manylinux2010_x86_64.whl (86.3MB)
[K     |████████████████████████████████| 86.3MB 32.7MB/s 
[?25hRequirement already satisfied, skipping upgrade: termcolor>=1.1.0 in /usr/local/lib/python3.6/dist-packages (from tensorflow==2.0) (1.1.0)
Requirement already satisfied, skipping upgrade: keras-applications>=1.0.8 in /usr/local/lib/python3.6/dist-packages (from tensorflow==2.0) (1.0.8)
Collecting tensorboard<2.1.0,>=2.0.0 (from tensorflow==2.0)
[?25l  Downloading https://files.pythonhosted.org/packages/9b/a6/e8ffa4e2ddb216449d34cfcb825ebb38206bee5c4553d69e7bc8bc2c5d64/tensorboard-2.0.0-py3-none-any.whl (3.8MB)
[K     |████████████████████████████████| 3.8MB 45.1MB/s 
[?25hRequirement already satisfied, skipping upgrade: six>=1.10.0 in /usr/local/lib/python3.6/dist-packages (from tensorflow==2.0) (1.12.0)
R
```

```python
import tensorflow as tf
import tensorflow.keras as keras
import numpy as np
from tensorflow.keras.layers import Dense, Dropout, Activation

np.random.seed(0)
```

```python
(x_train, y_train), (x_test, y_test) = keras.datasets.boston_housing.load_data()
```

```python
x_train.shape
```

Output:
```
(404, 13)
```

```python
y_train.shape
```

Output:
```
(404,)
```

```python
y_train.mean(), y_train.std()
```

Output:
```
(22.395049504950492, 9.199035423364862)
```

The keras model consists of multiple parts:

1. Construct the model layers, neurons per layer, activation function
1. Determine the loss function, metrics and optimization method
1. Fit the model to some data
1. Evaluate the model using the same metric

Some relevant documentation:
 - [initializers](https://keras.io/initializers/)
 - [loss functions](https://www.tensorflow.org/api_docs/python/tf/keras/losses)
 - [regularizations](https://keras.io/regularizers/)
 - [optimizers](https://keras.io/optimizers/)
 - [metrics](https://www.tensorflow.org/api_docs/python/tf/keras/metrics)


First, to construct a model, use the [Sequential](https://keras.io/getting-started/sequential-model-guide/) object. You can pass multiple layers to the sequential object. For this next step, we will only be using the [Dense](https://keras.io/layers/core/#dense) layers.

<!-- grade_id: cell-51879de7032d8823, points: 0, grade: false, solution: true -->
```python
# Create 3 hidden layers and an output layer with 10 neurons each and set them to the variable layers
# Use any activation you like such as "relu" or "tanh", you can alternate for each layer
# For your first run, try using the linear activation and then see if modifying the activations improved the result.
# Optional - give each layer a name and see how that shows up in the summary.
# YOUR CODE HERE

layers= [Dense(10, input_shape=(13,),
    activation='relu'),
    Dense(10, activation='relu'),
    Dense(10,activation='tanh'),
    Dense(1)]
model = keras.Sequential(layers)
#model.add(Activation('relu'))
model.summary()
```

Output:
```
Model: "sequential_1"
_________________________________________________________________
Layer (type)                 Output Shape              Param #   
=================================================================
dense_4 (Dense)              (None, 10)                140       
_________________________________________________________________
dense_5 (Dense)              (None, 10)                110       
_________________________________________________________________
dense_6 (Dense)              (None, 10)                110       
_________________________________________________________________
dense_7 (Dense)              (None, 1)                 11        
=================================================================
Total params: 371
Trainable params: 371
Non-trainable params: 0
_________________________________________________________________
```

<!-- grade_id: cell-ee74eff075f222ed, points: 5, grade: true, solution: false -->
```python
assert len(layers) == 4
for i,layer in enumerate(layers):
    assert isinstance(layers[i], keras.layers.Dense) 
    assert layer.weights[1].shape == [10,10,10,1][i]
```

<!-- grade_id: cell-9c5dcca57115d02a, points: 0, grade: false, solution: true -->
```python
# Set up the model to do the following:
# - use stochastic gradient descent to fit the model
# - use mean absolute error as its loss function
# - use mean absolute error as one of its metrics
# - use mean squared error as one of its metrics
# YOUR CODE HERE

sgd = keras.optimizers.SGD(lr=0.01, decay=1e-6, momentum=0.9, nesterov=True)
model.compile(loss='mean_absolute_error',
              optimizer=sgd,
              metrics=['mean_absolute_error','mean_squared_error'])
```

<!-- grade_id: cell-0746e29c34118826, points: 5, grade: true, solution: false -->
```python
assert isinstance(model.optimizer, keras.optimizers.SGD)
assert model.loss in ["mae", "mean_absolute_error"]
assert "mae" in model.metrics or "mean_absolute_error" in model.metrics
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-35-7f54c4b0d3e5> in <module>()
      1 assert isinstance(model.optimizer, keras.optimizers.SGD)
      2 assert model.loss in ["mae", "mean_absolute_error"]
----> 3 assert "mae" in model.metrics or "mean_absolute_error" in model.metrics

AssertionError:
```

```python
# Now fit the model
model.fit(x_train, y_train, epochs=50)
```

```python
model.summary()
```

Output:
```
Model: "sequential_6"
_________________________________________________________________
Layer (type)                 Output Shape              Param #   
=================================================================
dense_26 (Dense)             (None, 10)                140       
_________________________________________________________________
dense_27 (Dense)             (None, 10)                110       
_________________________________________________________________
dense_28 (Dense)             (None, 10)                110       
_________________________________________________________________
dense_29 (Dense)             (None, 1)                 11        
=================================================================
Total params: 371
Trainable params: 371
Non-trainable params: 0
_________________________________________________________________
```

```python
# Here we can evaluate how our model does based on the test data
model.evaluate(x_test, y_test)
```

Output:
```
102/1 [=================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================
```

Now let's try another optimizer instead of stochastic gradient descent. [Adam](https://keras.io/optimizers/#adam) is the recommended default for training neural networks since it usually performs quite well. In the next cell, compile the model with adam instead of sgd and use the same loss and metrics. After compiling, fit the data for as many epochs as you think it takes to see the value start to converge.

<!-- grade_id: cell-bee688f3d919be61, points: 0, grade: false, solution: true -->
```python
# Compile the model using adam
# YOUR CODE HERE
layers= [Dense(10, input_shape=(13,),
    activation='relu'),
    Dense(10, activation='relu'),
    Dense(10,activation='tanh'),
    Dense(1)]
model = keras.Sequential(layers)
ad=keras.optimizers.Adam(learning_rate=0.001, beta_1=0.9, beta_2=0.999, amsgrad=False)
model.compile(loss='mean_absolute_error',
              optimizer=ad,
              metrics=['mean_absolute_error','mean_squared_error'])
model.fit(x_train, y_train, epochs=50)
```

```python
model.summary()
```

Output:
```
Model: "sequential_8"
_________________________________________________________________
Layer (type)                 Output Shape              Param #   
=================================================================
dense_34 (Dense)             (None, 10)                140       
_________________________________________________________________
dense_35 (Dense)             (None, 10)                110       
_________________________________________________________________
dense_36 (Dense)             (None, 10)                110       
_________________________________________________________________
dense_37 (Dense)             (None, 1)                 11        
=================================================================
Total params: 371
Trainable params: 371
Non-trainable params: 0
_________________________________________________________________
```

<!-- grade_id: cell-fe8cd91b8b780964, points: 5, grade: true, solution: false -->
```python
assert isinstance(model.optimizer, keras.optimizers.Adam)
assert model.loss in ["mae", "mean_absolute_error"]
assert "mae" in model.metrics or "mean_absolute_error" in model.metrics
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-69-c7e862517747> in <module>()
      1 assert isinstance(model.optimizer, keras.optimizers.Adam)
      2 assert model.loss in ["mae", "mean_absolute_error"]
----> 3 assert "mae" in model.metrics or "mean_absolute_error" in model.metrics

AssertionError:
```

```python
model.evaluate(x_test, y_test)
```

Output:
```
102/1 [=================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================
```

Now recreate the model with dropout layers. Add 2 dropout layers, each after the first or second layer respectively. Select a low value of dropout that results in a good score.

<!-- grade_id: cell-d5040e44398f8dc9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
layers= [Dense(10, input_shape=(13,), activation='relu'),Dropout(.015),
        Dense(10, activation='relu'), Dropout(.015),
        Dense(10,activation='relu'), 
        Dense(1)]
model = keras.Sequential(layers)
```

<!-- grade_id: cell-b6c7e6e2ddc6742a, points: 5, grade: true, solution: false -->
```python
assert len(layers) == 6
assert isinstance(layers[1], Dropout)
assert isinstance(layers[3], Dropout)
for i,layer in enumerate(layers):
    if i not in [1,3]:
        assert isinstance(layers[i], keras.layers.Dense) 
        assert layer.weights[1].shape == [10,0,10,0,10,1][i]
```

```python
model.compile(optimizer='adam', loss='mae', metrics=['mae', "mse"])
model.fit(x_train, y_train, epochs=500, verbose=0)
model.summary()
```

Output:
```
Model: "sequential_17"
_________________________________________________________________
Layer (type)                 Output Shape              Param #   
=================================================================
dense_80 (Dense)             (None, 10)                140       
_________________________________________________________________
dropout_16 (Dropout)         (None, 10)                0         
_________________________________________________________________
dense_81 (Dense)             (None, 10)                110       
_________________________________________________________________
dropout_17 (Dropout)         (None, 10)                0         
_________________________________________________________________
dense_82 (Dense)             (None, 10)                110       
_________________________________________________________________
dense_83 (Dense)             (None, 1)                 11        
=====================================================
```

```python
model.evaluate(x_test, y_test)
```

Output:
```
102/1 [=================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================
```

Select a dropout rate that gets an okay score.

<!-- grade_id: cell-538bede4947fe364, points: 0, grade: true, solution: false -->
```python
assert model.evaluate(x_test, y_test)[0] < 7
```

Output:
```
102/1 [=================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================================
```

## Feedback

<!-- grade_id: feedback, points: 0, grade: false, solution: true -->
```python
def feedback():
    """Provide feedback on the contents of this exercise
    
    Returns:
        string
    """
    # YOUR CODE HERE
    return "Asserts are broken"
```

<!-- grade_id: feedback-tests, points: 0, grade: true, solution: false -->
```python

```